<a href="https://colab.research.google.com/github/phreakyphoenix/MXNet-GluonCV-AWS-Coursera/blob/master/Module_5_Assignment_LeNet_on_MNIST.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Graded Assessment

In this assessment you will write a full end-to-end training process using gluon and MXNet. We will train the LeNet-5 classifier network on the MNIST dataset. The network will be defined for you but you have to fill in code to prepare the dataset, train the network, and evaluate it's performance on a held out dataset.

```python
#Check CUDA version
!nvcc --version
```

Output:
```
nvcc: NVIDIA (R) Cuda compiler driver
Copyright (c) 2005-2019 NVIDIA Corporation
Built on Sun_Jul_28_19:07:16_PDT_2019
Cuda compilation tools, release 10.1, V10.1.243
```

```python
#Install appropriate MXNet version
'''
For eg if CUDA version is 10.0 choose mxnet cu100mkl 
where cu adds CUDA GPU support
and mkl adds Intel CPU Math Kernel Library support
'''
!pip install mxnet-cu101mkl gluoncv
```

Output:
```
Collecting mxnet-cu101mkl
[?25l  Downloading https://files.pythonhosted.org/packages/3d/4b/e51dc49ca5fe6564028e7c91b10a3f79c00d710dd691b408c77597df5883/mxnet_cu101mkl-1.6.0-py2.py3-none-manylinux1_x86_64.whl (711.0MB)
[K     |████████████████████████████████| 711.0MB 26kB/s 
[?25hCollecting gluoncv
[?25l  Downloading https://files.pythonhosted.org/packages/69/4d/d9d6b9261af8f7251977bb97be669a3908f72bdec9d3597e527712d384c2/gluoncv-0.6.0-py2.py3-none-any.whl (693kB)
[K     |████████████████████████████████| 696kB 43.7MB/s 
[?25hRequirement already satisfied: requests<3,>=2.20.0 in /usr/local/lib/python3.6/dist-packages (from mxnet-cu101mkl) (2.21.0)
Requirement already satisfied: numpy<2.0.0,>1.16.0 in /usr/local/lib/python3.6/dist-packages (from mxnet-cu101mkl) (1.18.2)
Collecting graphviz<0.9.0,>=0.8.1
  Downloading https://files.pythonhosted.org/packages/53/39/4ab213673844e0c004bed8a0781a0721a3f6bb23eb8854ee75c236428892/graphviz-0.8.4-py2.py3-none-any.whl
Collecting portalocker

```

```python
from pathlib import Path
from mxnet import gluon, metric, autograd, init, nd
import os
import mxnet as mx
```

```python
#I downloaded the files from Coursera and hosted on my gdrive:
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Go to this URL in a browser: https://accounts.google.com/o/oauth2/auth?client_id=947318989803-6bn6qk8qdgf4n4g3pfee6491hc0brc4i.apps.googleusercontent.com&redirect_uri=urn%3aietf%3awg%3aoauth%3a2.0%3aoob&response_type=code&scope=email%20https%3a%2f%2fwww.googleapis.com%2fauth%2fdocs.test%20https%3a%2f%2fwww.googleapis.com%2fauth%2fdrive%20https%3a%2f%2fwww.googleapis.com%2fauth%2fdrive.photos.readonly%20https%3a%2f%2fwww.googleapis.com%2fauth%2fpeopleapi.readonly

Enter your authorization code:
··········
Mounted at /content/drive
```

```python
# M5_DATA = Path(os.getenv('DATA_DIR', '../../data'), 'module_5')
M5_DATA = Path('/content/drive/My Drive/CourseraWork/MXNetAWS/data/module_5')
M5_IMAGES = Path(M5_DATA, 'images')
```

---
## Question 1

### Prepare and the data and construct the dataloader

* First, get the MNIST dataset from `gluon.data.vision.datasets`. Use
* Don't forget the ToTensor and normalize Transformations. Use `0.13` and `0.31` as the mean and standard deviation respectively
* Construct the dataloader with the batch size provide. Ensure that the train_dataloader is shuffled.

<font color='red'>**CAUTION!**</font>: Although the notebook interface has internet connectivity, the **autograders are not permitted to access the internet**. We have already downloaded the correct models and data for you to use so you don't need access to the internet. Set the `root` parameter to `M5_IMAGES` when using a preset dataset. Usually, in the real world, you have internet access, so setting the `root` parameter isn't required (and it's set to `~/.mxnet` by default).

<!-- grade_id: dataloader, points: 0, grade: false, solution: true -->
```python
import os
from pathlib import Path
from mxnet.gluon.data.vision import transforms
import numpy as np
def get_mnist_data(batch=128):
    """
    Should construct a dataloader with the MNIST Dataset with the necessary transforms applied.
    
    :param batch: batch size for the DataLoader.
    :type batch: int
    
    :return: a tuple of the training and validation DataLoaders
    :rtype: (gluon.data.DataLoader, gluon.data.DataLoader)
    """
    
    def transformer(data, label):
        data = data.flatten().expand_dims(0).astype(np.float32)/255
        data = data-0.13/0.31
        label = label.astype(np.float32)
        return data, label

    train_dataset = gluon.data.vision.datasets.MNIST(root=M5_IMAGES, train=True, transform=transformer)
    validation_dataset = gluon.data.vision.datasets.MNIST(root=M5_IMAGES, train=False, transform=transformer)
    train_dataloader = gluon.data.DataLoader(train_dataset, batch_size=batch, last_batch='keep',shuffle=True)
    validation_dataloader = gluon.data.DataLoader(validation_dataset, batch_size=batch, last_batch='keep')
    
    return train_dataloader, validation_dataloader
```

<!-- grade_id: cell-3ec320323b376193, points: 3, grade: true, solution: false -->
```python
t, v = get_mnist_data()
assert isinstance(t, gluon.data.DataLoader)
assert isinstance(v, gluon.data.DataLoader)

d, l = next(iter(t))
assert d.shape == (128, 1, 28, 28) #check Channel First and Batch Size
assert l.shape == (128,)

assert nd.max(d).asscalar() <= 2.9 # check for normalization
assert nd.min(d).asscalar() >= -0.5 # check for normalization
```

---

## Question 2

### Write the training loop

* Create the loss function. This should be a loss function suitable for multi-class classification.
* Create the metric accumulator. This should the compute and store the accuracy of the model during training
* Create the trainer with the `adam` optimizer and learning rate of `0.002`
* Write the training loop

<!-- grade_id: train, points: 0, grade: false, solution: true -->
```python
def train(network, training_dataloader, batch_size, epochs):
    """
    Should take an initialized network and train that network using data from the data loader.
    
    :param network: initialized gluon network to be trained
    :type network: gluon.Block
    
    :param training_dataloader: the training DataLoader provides batches for data for every iteration
    :type training_dataloader: gluon.data.DataLoader
    
    :param batch_size: batch size for the DataLoader.
    :type batch_size: int
    
    :param epochs: number of epochs to train the DataLoader
    :type epochs: int
    
    :return: tuple of trained network and the final training accuracy
    :rtype: (gluon.Block, float)
    """
    trainer = gluon.Trainer(network.collect_params(), 'adam',
                            {'learning_rate': 0.002})
    metric = mx.metric.Accuracy()
    
    for epoch in range(epochs):
        train_loss =0.
        for data,label in training_dataloader:
        
#             print (data.shape)
#             print (label.shape)
            with autograd.record():
                output = network(data)
                loss=mx.ndarray.softmax_cross_entropy(output,label)
            loss.backward()

            trainer.step(batch_size)
            train_loss += loss.mean().asscalar()
            metric.update(label, output)
            
        print (epoch , metric.get()[1])    
        training_accuracy = metric.get()[1]
    return network, training_accuracy
```

Let's define and initialize a network to test the train function.

```python
net = gluon.nn.Sequential()
net.add(gluon.nn.Conv2D(channels=6, kernel_size=5, activation='relu'),
        gluon.nn.MaxPool2D(pool_size=2, strides=2),
        gluon.nn.Conv2D(channels=16, kernel_size=3, activation='relu'),
        gluon.nn.MaxPool2D(pool_size=2, strides=2),
        gluon.nn.Flatten(),
        gluon.nn.Dense(120, activation="relu"),
        gluon.nn.Dense(84, activation="relu"),
        gluon.nn.Dense(10))
net.initialize(init=init.Xavier())
```

<!-- grade_id: cell-9d116b0901b24fcc, points: 5, grade: true, solution: false -->
```python
n, ta = train(net, t, 128, 5)
assert ta >= .95

d, l = next(iter(v))
p = (n(d).argmax(axis=1))
assert (p.asnumpy() == l.asnumpy()).sum()/128.0 > .95
```

Output:
```
0 0.93415
1 0.9572583333333333
2 0.9668111111111111
3 0.972375
4 0.97606
```

---
## Question 3

### Write the validation loop

* Create the metric accumulator. This should the compute and store the accuracy of the model on the validation set
* Write the validation loop

<!-- grade_id: sum_of_squares_application, points: 0, grade: false, solution: true -->
```python
def validate(network, validation_dataloader):
    """
    Should compute the accuracy of the network on the validation set.
    
    :param network: initialized gluon network to be trained
    :type network: gluon.Block
    
    :param validation_dataloader: the training DataLoader provides batches for data for every iteration
    :type validation_dataloader: gluon.data.DataLoader
    
    :return: validation accuracy
    :rtype: float
    """
    val_acc = mx.metric.Accuracy()
    for data,label in validation_dataloader:
        output = network(data)
        val_acc.update(label,output)
    print (val_acc.get()[1])
    return val_acc.get()[1]
```

<!-- grade_id: cell-5a2130f381810258, points: 2, grade: true, solution: false -->
```python
assert validate(n, v) > .95
```

Output:
```
0.9896
```

```python

```